```python
# import and initialize otter
import otter
grader = otter.Notebook("p1.ipynb")
```

```python
import public_tests
```

```python
# PLEASE FILL IN THE DETAILS

# project: p1
# submitter: NETID1
# hours: ????
```

## Project 1: Introduction

**Question 1:** "Hello, World!" program.

```python
name = "World"
greeting = "Hello, " + name + "!"

greeting
```

```python
grader.check("q1")
```

**Question 2:** Compute the current year.

```python
curr_year = 45 * 45 - 2

curr_year
```

```python
grader.check("q2")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

**SUBMISSION INSTRUCTIONS**:
 1. **Upload** the zipfile to Gradescope.
 2. Check **Gradescope** results as soon as the auto-grader execution gets completed.
 3. Your **final score** for this project is the score that you see on **Gradescope**.
 4. You are **allowed** to resubmit on Gradescope as many times as you want to.
 5. **Contact** a TA/PM if you lose any points on Gradescope for any **reasons**.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: 12.5)
```
>>> public_tests.check('q1', greeting)
All test cases passed!
```


Tests for q2 (points: 12.5)
```
>>> public_tests.check('q2', curr_year)
All test cases passed!
```
