```python
# import and initialize otter
import otter
grader = otter.Notebook("p6.ipynb")
```

```python
import public_tests
```

```python
# PLEASE FILL IN THE DETAILS
# enter none if you don't have a project partner
# you will have to add your partner as a group member on Gradescope even after you fill this

# project: p6
# submitter: alghazzy
# partner: jmbruggink
# hours: 4
```

# Project 6: Power Generators in Wisconsin

## Learning Objectives:

In this project, you will demonstrate how to:

* access and utilize data in CSV files,
* process real world datasets,
* use string methods and sorting function / method to order data.

Please go through [Lab-P6](https://git.doit.wisc.edu/cdis/cs/courses/cs220/cs220-f23-projects/-/tree/main/lab-p6) before working on this project. The lab introduces some useful techniques related to this project.

## Testing your code:

Along with this notebook, you must have downloaded the file `public_tests.py`. If you are curious about how we test your code, you can explore this file, and specifically the function `get_expected_json`, to understand the expected answers to the questions.

## Project Description:

In this project, we will use Python to analyze Electric Power Generation within the state of Wisconsin. The data used in this project has been adapted from [this dataset](https://www.eia.gov/electricity/data/eia860m/) maintained by the **US Energy Information Administration**. `power_generators.csv` has data of all 613 power generators within the state of Wisconsin that were in operation as of August 2023. This file includes information about the capacities, locations, and the technologies used by the generators.

## Dataset:

A small portion of the dataset `power_generators.csv` you will be working with for this project is reproduced here:

entity_id|entity_name|plant_id|plant_name|generator_id|county|net_summer_capacity|net_winter_capacity|technology|latitude|longitude
---|---|---|---|---|---|---|---|---|---|---
13781|Northern States Power Co - Minnesota|1756|Saxon Falls|1|Iron|0.5|0.5|Conventional Hydroelectric|46.5392|-90.3742
13781|Northern States Power Co - Minnesota|1756|Saxon Falls|2|Iron|0.5|0.6|Conventional Hydroelectric|46.5392|-90.3742
20847|Wisconsin Electric Power Co|1775|Brule|1|Florence|1.3|1.3|Conventional Hydroelectric|45.9472|-88.2189
20847|Wisconsin Electric Power Co|1775|Brule|2|Florence|2|2|Conventional Hydroelectric|45.9472|-88.2189
20847|Wisconsin Electric Power Co|1775|Brule|3|Florence|2|2|Conventional Hydroelectric|45.9472|-88.2189

Each row of data represents a **single** generator. The columns contain the following data about each generator (along with the correct data type you **must** represent it as):

1. `entity_id` - the **ID** of the **entity** that operates the Power Generator (`int`)
2. `entity_name` - the **name** of the **entity** that operates the Power Generator (`str`)
3. `plant_id` - the **ID** of the **Power Plant** hosting the Power Generator (`int`)
4. `plant_name` - the **name** of the **Power Plant** hosting the Power Generator (`str`)
5. `generator_id` - the **ID** of the specific **Power Generator** within its Power Plant (`str`)
6. `county` - the **name** of the **county** that the **Power Plant** is located in (`str`)
7. `net_summer_capacity` - the maximum **capacity** of the **Power Generator** (in units of MW) during the Summer months (`float`)
8. `net_winter_capacity` - the maximum **capacity** of the **Power Generator** (in units of MW) during the Winter months (`float`)
9. `technology` - the **technology** used by the **Power Generator** (`str`)
10. `latitude` - the **latitude** where the **Power Plant** is located (`float`)
11. `longitude` - the **longitude** where the **Power Plant** is located (`float`)

You can find more details on the dataset in [Lab-P6](https://git.doit.wisc.edu/cdis/cs/courses/cs220/cs220-f23-projects/-/tree/main/lab-p6).

## Project Requirements:

You **may not** hardcode indices in your code unless specified in the question. If you are not sure what hardcoding is, here is a simple test you can use to determine whether you have hardcoded:

*If we were to change the data (e.g. add more power generators, remove some power generators, or swap some columns or rows), would your code still find the correct answer to the question as it is asked?*

If your answer to that question is *No*, then you have likely hardcoded something. Please reach out to TAs/PMs during office hours to find out how you can **avoid hardcoding**.

**Store** your final answer for each question in the **variable specified for each question**. This step is important because Otter grades your work by comparing the value of this variable against the correct answer.

For some of the questions, we'll ask you to write (then use) a function to compute the answer.  If you compute the answer **without** creating the function we ask you to write, the Gradescope autograder will **deduct** points, even if the way you did it produced the correct answer.

Required Functions:
- `process_csv`
- `cell`
- `find_entities_with_phrase`
- `num_generators_by`
- `find_indices_within`
- `median`
- `total_summer_capacity_of`
- `avg_winter_capacity_of`
    
Students are only allowed to use Python commands and concepts that have been taught in the course prior to the release of P6. Therefore, **you should not use concepts/modules such as dictionaries, or the pandas module, to name a few examples**. Otherwise, the Gradescope autograder will **deduct** points, even if the way you did it produced the correct answer.

For more details on what will cause you to lose points during code review and specific requirements, please take a look at the [Grading rubric](https://git.doit.wisc.edu/cdis/cs/courses/cs220/cs220-f23-projects/-/blob/main/p6/rubric.md).

<h2 style="color:red">Warning (Note on Academic Misconduct):</h2>

Under any circumstances, **no more than two students are allowed to work together on a project** as mentioned in the course policies. If your code is flagged by our code similarity detection tools, **both partners will be responsible** for sharing/copying the code, even if the code is shared/copied by one of the partners with/from other non-partner student(s). Note that each case of plagiarism will be reported to the Dean of Students with a zero grade on the project. **If you think that someone cannot be your project partner then don’t make that student your lab partner.** 
<font color = "red">Project partners must submit only one copy of their project on Gradescope, but they must include the names of both partners.</font>

## Questions and Functions:

Let us start by importing all the modules we will need for this project.

```python
# it is considered a good coding practice to place all import statements at the top of the notebook
# please place all your import statements in this cell if you need to import any more modules for this project
import csv
```

### Function 0: `cell`

Copy and paste the `process_csv` and `cell` functions from your Lab-P6 notebook to the cell below.

You are expected to call the `process_csv` function correctly, and read the data on `power_generators.csv`. After reading the file, define the `csv_header`, and `csv_rows` variables as in Lab-P6, and define the `cell` function.

**Important:** You **must** only use the `cell` function to extract data from the dataset. If you extract any data without explicitly using this function, the Gradescope autograder will **deduct points**.

```python
def process_csv(filename):
    example_file = open(filename, encoding="utf-8")
    example_reader = csv.reader(example_file)
    example_data = list(example_reader)
    example_file.close()
    return example_data

def cell(row_idx, col_name):
    col_idx = csv_header.index(col_name)
    val = csv_rows[row_idx][col_idx]
    if val == " ":
        return None
    elif col_name == 'entity_id':
        val = int(val)
    elif col_name == 'net_summer_capacity':
        val = float(val)
    elif col_name == 'plant_id':
        val = int(val)
    elif col_name == 'plant_name':
        val = str(val)
    elif col_name == 'generator_id':
        val = str(val)
    elif col_name == 'county':
        val = str(val)
    elif col_name == 'net_winter_capacity':
        val = float(val)
    elif col_name == 'technology':
        val = str(val)
    elif col_name == 'latitude':
        val = float(val)
    elif col_name == 'longitude':
        val = float(val)
    return val

csv_data = process_csv("power_generators.csv")
csv_header = csv_data[0]
csv_rows = csv_data[1:]
num_rows = len(csv_rows)
```

```python
grader.check("cell")
```

Output:
```
cell results: All test cases passed!
cell - 1 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'variables `csv_data`, `csv_header`, and `csv_rows` are not defined as expected (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
cell - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'function does not return `None` when data is missing (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
cell - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'function does not typecast to `int` or `float` when expected (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
cell - 4 message: Note that the Gradescope autograder will deduct point
```

**Question 1:** What **unique** technologies (`technology`) are used by the power generators in Wisconsin?

Your output **must** be a *list* which stores all the **unique** technologies (i.e., without any duplicates). The order **does not** matter.

```python
# compute and store the answer in the variable 'technologies', then display it
technologies = []
list_with_duplicates = []
for idx in range(num_rows):
    list_with_duplicates.append(cell(idx, 'technology'))
for value in list_with_duplicates:
    if value not in technologies:
        technologies.append(value)
list_without_duplicates = []
for value in technologies:
    list_without_duplicates.append(value)

technologies = list_without_duplicates
technologies
```

Output:
```
['Conventional Hydroelectric',
 'Petroleum Liquids',
 'Wood/Wood Waste Biomass',
 'Natural Gas Steam Turbine',
 'Natural Gas Fired Combustion Turbine',
 'Natural Gas Fired Combined Cycle',
 'Conventional Steam Coal',
 'Nuclear',
 'Natural Gas Internal Combustion Engine',
 'Petroleum Coke',
 'Onshore Wind Turbine',
 'Landfill Gas',
 'Other Waste Biomass',
 'Solar Photovoltaic',
 'Batteries']
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
q1 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'not all rows are looped through (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q1 - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q1 - 4 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`cell` function is not used to read data (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

**Question 2:** How many power generators are in the `county` *Dane*?

```python
# compute and store the answer in the variable 'count_dane', then display it
count_dane = 0
for idx in range(num_rows):  
    if cell(idx, 'county') == 'Dane': 
        count_dane += 1
count_dane
```

Output:
```
47
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
q2 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'not all rows are looped through (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q2 - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q2 - 4 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`cell` function is not used to read data (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

**Question 3:** What is the **total** `net_summer_capacity` of all the power generators in Wisconsin?

Your answer **must** be a **float** that represents the total `net_summer_capacity`. You **must** **ignore** all power generators whose `net_summer_capacity` data is **missing**.

```python
# compute and store the answer in the variable 'total_summer_capacity', then display it
total_summer_capacity = 0.0
for idx in range(num_rows):
    cap = 0
    if type(cell(idx, 'net_summer_capacity')) == float:
        cap = cell(idx,'net_summer_capacity')
        total_summer_capacity += cap
    else:
        continue
total_summer_capacity
```

Output:
```
16831.099999999944
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
q3 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'not all rows are looped through (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q3 - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q3 - 4 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`cell` function is not used to read data (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

### Function 1: `find_entities_with_phrase(phrase)`

We require you to complete the below function. You can review string methods from any of these lecture slides: [Mike](https://canvas.wisc.edu/courses/374263/files/folder/Mikes_Lecture_Notes/Lec13_Strings), [Gurmail](https://git.doit.wisc.edu/cdis/cs/courses/cs220/cs220-lecture-material/-/tree/main/f23/Gurmail_Lecture_Notes/13_Strings), or [Cole](https://git.doit.wisc.edu/cdis/cs/courses/cs220/cs220-lecture-material/-/tree/main/f23/Cole_Lecture_Notes/13_Strings).

```python
def find_entities_with_phrase(phrase):
    list1 = []
    for idx in range(len(csv_rows)):
        scanner = cell(idx, 'entity_name').split()
        for words in scanner:
            if words.lower() == phrase.lower():
                list1.append(cell(idx, 'entity_name'))               
    return set(list1)
```

```python
grader.check("find_entities_with_phrase")
```

Output:
```
find_entities_with_phrase results: All test cases passed!
find_entities_with_phrase - 1 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'not all rows are looped through (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
find_entities_with_phrase - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'case insensitive matches are not considered (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
find_entities_with_phrase - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'duplicates are not removed (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
find_entities_with_phrase - 4 message: Note that the Gradescope autograder wi
```

**Question 4:** Find all entity names (`entity_name`) that contain the string *"Madison"* (case insensitive).
    
Your output **must** be a **list**. The order **does not** matter. You **must** use the `find_entities_with_phrase` function to answer this question.

```python
# compute and store the answer in the variable 'madison_entities', then display it

madison_entities = find_entities_with_phrase("Madison")
madison_entities = list(madison_entities)
madison_entities
```

Output:
```
['Madison Gas & Electric Co']
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
q4 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`find_entities_with_phrase` function is not used to answer (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

**Question 5:** Find all entity names (`entity_name`) that contain **either** *"Wisconsin"* **or** *"Power"* (case insensitive).

If an entity's name contains **both** *"Wisconsin"* and *"Power"*, then the `entity_name` must be included **only once** in your list.

Your output **must** be a **list**. The order **does not** matter.

**Hint**: You can use the `find_entities_with_phrase` function on *"Wisconsin"* and *"Power"* to answer this question.

```python
entities_contain_wisconsin_power = []
new_list = list(find_entities_with_phrase("Wisconsin")) + list(find_entities_with_phrase("Power"))

for entity_name in new_list:
    entity_lower = entity_name.lower()  
    if "wisconsin" in entity_lower or "power" in entity_lower:
        if entity_lower not in [e.lower() for e in entities_contain_wisconsin_power]:
            entities_contain_wisconsin_power.append(entity_name)
entities_contain_wisconsin_power
```

Output:
```
['Wisconsin River Power Company',
 'Northwestern Wisconsin Elec Co',
 'Wisconsin Power & Light Co',
 'State of Wisconsin',
 'Wisconsin Electric Power Co',
 'Wisconsin Public Service Corp',
 'HQC Rock River Solar Power Generation Station LLC',
 'Dahlberg Light & Power Co',
 'Dairyland Power Coop',
 'Northern States Power Co - Minnesota',
 'North Central Power Co Inc',
 'Consolidated Water Power Co']
```

```python
grader.check("q5")
```

Output:
```
q5 results: All test cases passed!
q5 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q5 - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`cell` function is not used to read data (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

**Question 6:** Find all entity names (`entity_name`) that contain **both** *"Solar"* **and** *"LLC"* (case insensitive).

Your output **must** be a **list**. The order **does not** matter.

**Hint**: You can use the `find_entities_with_phrase` function on *"Solar"* and *"LLC"* to answer this question.

```python
# compute and store the answer in the variable 'entities_contain_solar_llc', then display it
entities_contain_solar_llc = []
new_list = list(find_entities_with_phrase("Solar")) + list(find_entities_with_phrase("LLC"))
for entity_name in new_list:
    entity_lower = entity_name.lower()  
    if "solar" in entity_lower and "llc" in entity_lower:
        if entity_lower not in [e.lower() for e in entities_contain_solar_llc]:
            entities_contain_solar_llc.append(entity_name)
entities_contain_solar_llc
```

Output:
```
['Flambeau Solar Partners, LLC',
 'HQC Rock River Solar Power Generation Station LLC',
 'Dane County Solar LLC',
 'Middleton Biogas Solar, LLC',
 'Butter Solar, LLC']
```

```python
grader.check("q6")
```

Output:
```
q6 results: All test cases passed!
q6 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q6 - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`cell` function is not used to read data (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

```python

```

**Question 7:** Find the generator IDs (`generator_id`) of all the generators that use the `technology` *"Wood/Wood Waste Biomass"* within the power plant with the `plant_id` *50614*.

Your output **must** be a *list*. The IDs **must** be sorted in **descending (alphabetical) order**.

```python
# compute and store the answer in the variable 'plant_50614_generators', then display it
plant_50614_generators = []
for idx in range(len(csv_rows)):
    if cell(idx, 'technology') == "Wood/Wood Waste Biomass": 
        if cell(idx, 'plant_id') == 50614:
            plant_50614_generators.append(cell(idx, 'generator_id'))
plant_50614_generators.sort(reverse = True)
plant_50614_generators
```

Output:
```
['WEST', 'GEN1']
```

```python
grader.check("q7")
```

Output:
```
q7 results: All test cases passed!
q7 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'not all rows are looped through (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q7 - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'the output is not sorted correctly (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q7 - 4 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q7 - 5 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`cell` function is not used to read data (-1)'. The public te
```

**Question 8:**  What are the power plants (`plant_name`) that contain generators which use the `technology` *Conventional Hydroelectric* and have a `net_summer_capacity` greater than *5*?

You **must** **ignore** all generators with **missing** `net_summer_capacity` data.

Your output **must** be a *list* of **unique** plant names (`plant_name`). The names **must** be sorted in **ascending (alphabetical) order**.

```python
# compute and store the answer in the variable 'powerful_hydro_electric_plants', then display it
powerful_hydro_electric_plants = []
new_list = []
for idx in range(len(csv_rows)):
    if cell(idx, 'technology') == 'Conventional Hydroelectric':
        if type(cell(idx, 'net_summer_capacity')) == float and cell(idx, 'net_summer_capacity') > 5.0:
            powerful_hydro_electric_plants.append(cell(idx, 'plant_name'))
    else:
        continue
for value in powerful_hydro_electric_plants:
    if value not in new_list:
        new_list.append(value)
list_without_duplicates = []
for value in new_list:
    list_without_duplicates.append(value)
powerful_hydro_electric_plants = list_without_duplicates
powerful_hydro_electric_plants.sort()
powerful_hydro_electric_plants    
```

Output:
```
['Cornell',
 'Flambeau Hydroelectric Station',
 'Grandfather Falls',
 'Holcombe',
 'Jim Falls',
 'Kilbourn',
 'Prairie Du Sac']
```

```python
grader.check("q8")
```

Output:
```
q8 results: All test cases passed!
q8 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'not all rows are looped through (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q8 - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'the output is not sorted correctly (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q8 - 4 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect comparison operators are used (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q8 - 5 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used (-1)'. The public tes
```

### Function 2: `num_generators_by(entity_name)`

We require you to complete the below function.

```python
def num_generators_by(entity_name):
    """
    num_generators_by(entity_name) returns the number of
    power generators operated by the given `entity_name`
    (case sensitive match).
    """
    num_generators = 0
    for idx in range(len(csv_rows)):
        if entity_name == cell(idx, "entity_name"):
            num_generators += 1      
    return num_generators
```

```python
grader.check("num_generators_by")
```

Output:
```
num_generators_by results: All test cases passed!
num_generators_by - 1 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'not all rows are looped through (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
num_generators_by - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'function logic is incorrect (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
num_generators_by - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`cell` function is not used to read data (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

**Question 9:** How **many** generators are operated by the entity (`entity_name`) *Madison Gas & Electric Co*?

You **must** use the `num_generators_by` function to answer this question.

```python
# compute and store the answer in the variable 'num_generators_by_mge', then display it
num_generators_by_mge = 0
num_generators_by_mge = num_generators_by("Madison Gas & Electric Co")
num_generators_by_mge
```

Output:
```
19
```

```python
grader.check("q9")
```

Output:
```
q9 results: All test cases passed!
q9 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`num_generators_by` function is not used to answer (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

**Question 10:** How **many** generators are operated by entities whose name (`entity_name`) **contains** the **phrase** *River* (case insensitive)?

You **must** use the `num_generators_by` and `find_entities_with_phrase` functions to answer this question.

```python
# compute and store the answer in the variable 'num_generators_by_river', then display it
num_generators_by_river = 0
new_list = find_entities_with_phrase('River')
for value in new_list:
    num_generators_by_river += num_generators_by(value)
num_generators_by_river
```

Output:
```
10
```

```python
grader.check("q10")
```

Output:
```
q10 results: All test cases passed!
q10 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q10 - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`num_generators_by` function is not used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q10 - 4 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`find_entities_with_phrase` function is not used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

**Question 11:** Which entity (`entity_name`) operates the **most** number of generators within Wisconsin?

You **must** use the `num_generators_by` function to answer this question. You do **not** have to worry about any ties. There is a **unique** entity with the most number of generators in the dataset.

**Hint**: You must first create a list of unique entities (`entity_name`) in the dataset, then loop through them to find the entity with the most number of generators.

**Extra Hint**: If you are clever about it, you can generate a list of all the unique entities using the `find_entities_with_phrase` function. You do **not** have to answer this question that way, but try to find out how you can use it here!

```python
# compute and store the answer in the variable 'most_generators_entity', then display it
most_generators_entity = ' '
most_generators_number = 0
entity_names = []
list_with_duplicates = []
for idx in range(num_rows):
    list_with_duplicates.append(cell(idx, 'entity_name'))
for value in list_with_duplicates:
    if value not in entity_names:
        entity_names.append(value)
list_without_duplicates = []
for value in entity_names:
    list_without_duplicates.append(value)
entity_names = list_without_duplicates
for value in entity_names:
    if num_generators_by(value) > most_generators_number:
        most_generators_number = num_generators_by(value)
        most_generators_entity = value
most_generators_entity
```

Output:
```
'Northern States Power Co - Minnesota'
```

```python
grader.check("q11")
```

Output:
```
q11 results: All test cases passed!
q11 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'all entities are looped through instead of just the unique entities (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q11 - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q11 - 4 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`num_generators_by` function is not used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

### Function 3: `find_indices_within(lat_min, lat_max, long_min, long_max)` 

We require you to complete the below function.

```python
def find_indices_within(lat_min, lat_max, long_min, long_max):
    """
    find_indices_within(lat_min, lat_max, long_min, long_max) returns a
    list of *row indices* of all generators located within the
    latitudes `lat_min` and `lat_max` (both inclusive) and the
    longitudes `long_min` and `long_max` (both inclusive).
    """
    new_list = []
    for idx in range(len(csv_rows)):
        if lat_min <= cell(idx, 'latitude') <= lat_max:
            if long_min <= cell(idx, 'longitude') <= long_max:
                new_list.append(idx)
    return new_list
```

```python
grader.check("find_indices_within")
```

Output:
```
find_indices_within results: All test cases passed!
find_indices_within - 1 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'not all rows are looped through (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
find_indices_within - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect comparison operators are used (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
find_indices_within - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'function logic is incorrect (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
find_indices_within - 4 message: Note that the Gradescope autograder will deduct points if your code fai
```

**Question 12:** How **many** power generators are located **within** the *City of Milwaukee* (`42.9870 <= latitude <= 43.1936`, `-88.0636 <= longitude <= -87.8727`)?

Note that simply checking if the `county` is *Milwaukee* will lead you to count generators that are within *Milwaukee County*, but not within the City. Use the coordinates given above to determine the generators that lie within the City.

You **must** use the `find_indices_within` function to answer this question.

```python
# compute and store the answer in the variable 'num_generators_in_milwaukee', then display it
new_list = find_indices_within(42.9870, 43.1936, -88.0636, -87.8727)
num_generators_in_milwaukee = len(new_list)
num_generators_in_milwaukee
```

Output:
```
8
```

```python
grader.check("q12")
```

Output:
```
q12 results: All test cases passed!
q12 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'length of the list is found by looping through it (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q12 - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`find_indices_within` function is not used to answer (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

**Question 13:** What are the **unique** technologies (`technology`) used by power generators located **near** the *University of Wisconsin-Madison Department of Computer Sciences* (`43.0675 <= latitude <= 43.0725`, `-89.4100 <= longitude <= -89.4000`)?

You may assume that any power generator that lies within the coordinates given above are **near** the *University of Wisconsin-Madison Department of Computer Sciences*. You **must** use the `find_indices_within` function to answer this question.

You **must** return a **list** of **unique** technologies. The order **does not** matter. 

```python
# compute and store the answer in the variable 'uw_madison_technologies', then display it
new_list = find_indices_within(43.0675, 43.0725, -89.4100, -89.4000)
    
uw_madison_technologies = []
list_with_duplicates = []
for idx in new_list:
    list_with_duplicates.append(cell(idx, 'technology'))
for value in list_with_duplicates:
    if value not in uw_madison_technologies:
        uw_madison_technologies.append(value)
list_without_duplicates = []
for value in uw_madison_technologies:
    list_without_duplicates.append(value)

uw_madison_technologies = list_without_duplicates
uw_madison_technologies
    
```

Output:
```
['Natural Gas Steam Turbine', 'Petroleum Liquids']
```

```python
grader.check("q13")
```

Output:
```
q13 results: All test cases passed!
q13 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'duplicates are not removed from the list (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q13 - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q13 - 4 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`find_indices_within` function is not used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

**Question 14:** Which power plant (`plant_name`) in *North Wisconsin* (`44.9657 <= latitude <= 46.6989`, `-92.1908 <= longitude <= -87.6449`) has the generator with the **highest** `net_summer_capacity`?

You may assume that any power generator that lies within the coordinates given above are **in** *North Wisconsin*. You **may** assume that **none** of the `net_summer_capacity` values of any of the power generators within this area are **missing**.

You do **not** have to worry about any ties. There is a **unique** generator in *North Wisconsin* with the highest `net_summer_capacity`.

You **must** use the `find_indices_within` function to answer this question.

```python
# compute and store the answer in the variable 'north_wisconsin_most_powerful', then display it
new_list = find_indices_within(44.9657, 46.6989, -92.1908, -87.6449)
north_wisconsin_most_powerful = ' '
highest_sum_capacity = 0
for idx in new_list:
    if cell(idx, 'net_summer_capacity') > highest_sum_capacity:
        highest_sum_capacity = cell(idx, 'net_summer_capacity')
        north_wisconsin_most_powerful = cell(idx, 'plant_name')
north_wisconsin_most_powerful
```

Output:
```
'West Marinette'
```

```python
grader.check("q14")
```

Output:
```
q14 results: All test cases passed!
q14 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q14 - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`find_indices_within` function is not used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q14 - 4 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`cell` function is not used to read data (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

### Function 4: `median(items)` 

We require you to complete the below function. You may **copy/paste** this function from your Lab-P6 notebook.

```python
def median(items):
    sorted_list = sorted(items)
    list_len = len(items)
    if list_len % 2 == 1:
        return sorted_list[list_len // 2]
    else:
        first_middle = sorted_list[list_len // 2] 
        second_middle = sorted_list[(list_len // 2) - 1]
        return (first_middle + second_middle) / 2
```

```python
grader.check("median")
```

Output:
```
median results: All test cases passed!
median - 1 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'function modifies the original input list (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
median - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'function logic is incorrect when inputs are of odd length (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
median - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'function logic is incorrect when inputs are of even length (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

**Question 15:** What is the **median** `net_winter_capacity` of *Conventional Hydroelectric* (`technology`) power generators **near** *Lake Winnebago* (`43.6961 <= latitude <= 44.3512`, `-88.5375 <= longitude <= -88.2713`)?

You may assume that any power generator that lies within the coordinates given above are **near** *Lake Winnebago*. You **may** assume that **none** of the `net_winter_capacity` values of any of the power generators within this area are **missing**.

You **must** use the `find_indices_within` and `median` functions to answer this question.

```python
# compute and store the answer in the variable 'winnebago_hydro_winter_capacity', then display 
winnebago_hydro_winter_capacity = 0
new_list = find_indices_within(43.6961, 44.3512, -88.5375, -88.2713)
neww_list = []
for idx in new_list:
    if cell(idx, 'technology') == 'Conventional Hydroelectric':
        neww_list.append(cell(idx, 'net_winter_capacity'))
winnebago_hydro_winter_capacity = median(neww_list)
winnebago_hydro_winter_capacity
```

Output:
```
0.7
```

```python
grader.check("q15")
```

Output:
```
q15 results: All test cases passed!
q15 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q15 - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`find_indices_within` function is not used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q15 - 4 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`median` function is not used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q15 - 5 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`cell` function is n
```

### Function 5: `total_summer_capacity_of(plant_name)` 

We **require** you to complete the below function. This function must take in a `plant_name` and return the **total** `net_summer_capacity` of all the power generators within that power plant. If a particular `plant_name` has **some** generators with **missing data**, those generators **must** be **ignored** while finding the sum.

This function can be **case-sensitive**. You **only** need to consider the power plants whose names **exactly match** `plant_name`.

```python
def total_summer_capacity_of(plant_name):
    """
    total_summer_capacity_of(plant_name) returns the total
    `net_summer_capacity` of all power generators within
    the given `plant_name`; this function **ignores**
    any generator whose `net_summer_capacity` data is
    missing.
    """
    var = 0
    for idx in range(len(csv_rows)):
        if cell(idx, 'plant_name') == plant_name:
            if type(cell(idx, 'net_summer_capacity')) == float:
                var += cell(idx, 'net_summer_capacity')
    return var
```

```python
grader.check("total_summer_capacity_of")
```

Output:
```
total_summer_capacity_of results: All test cases passed!
total_summer_capacity_of - 1 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'not all rows are looped through (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
total_summer_capacity_of - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'missing data is not ignored (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
total_summer_capacity_of - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'function logic is incorrect (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
total_summer_capacity_of - 4 message: Note that the Gradescope autograder will deduct points if 
```

**Question 16:** What is the **net summer capacity** of the **power plant** *Point Beach Nuclear Plant*?

The **net summer capacity** of a **power plant** refers to the **total** `net_summer_capacity` of **all** the generators within the power plant.

You **must** use the `total_summer_capacity_of` function to answer this question.

```python
# compute and store the answer in the variable 'point_beach_summer_capacity', then display it
point_beach_summer_capacity = total_summer_capacity_of('Point Beach Nuclear Plant')
point_beach_summer_capacity
```

Output:
```
1211.0
```

```python
grader.check("q16")
```

Output:
```
q16 results: All test cases passed!
q16 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`total_summer_capacity_of` function is not used to answer (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

**Question 17:** Find the **median** of the **net summer capacities** of **all** the **power plants** in Wisconsin.

The **net summer capacity** of a **power plant** refers to the **total** `net_summer_capacity` of **all** the generators within the power plant.

You **must** use the `total_summer_capacity_of` function to answer this question.

**WARNING**: You **must not** find the **median** across all the power **generators**. Multiple generators may belong to the same power plant. Instead, you **must** find the **median** across the **total** power generated by each **power plant**.

**Hint**: You must first make a list of all the **unique** power plants (`plant_name`) in the dataset, then make a **list** of all their **total** net summer capacities (using `total_summer_capacity_of`), and finally, find the **median** of this list.

```python
# compute and store the answer in the variable 'median_summer_capacity', then display it
median_summer_capacity = 0
plant_names = []
new_list = []
list_with_duplicates = []
for idx in range(num_rows):
    list_with_duplicates.append(cell(idx, 'plant_name'))
for value in list_with_duplicates:
    if value not in plant_names:
        plant_names.append(value)
list_without_duplicates = []
for value in plant_names:
    list_without_duplicates.append(value)

plant_names = list_without_duplicates
for value in plant_names:
    new_list.append(total_summer_capacity_of(value))
median_summer_capacity = median(new_list)
median_summer_capacity
```

Output:
```
7.0
```

```python
grader.check("q17")
```

Output:
```
q17 results: All test cases passed!
q17 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'all power plants are looped through instead of just the unique power plants (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q17 - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q17 - 4 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`total_summer_capacity_of` function is not used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q17 - 5 message: Note that the Gradescope autograder will deduct points if your code fails the follo
```

### Function 6: `avg_winter_capacity_of(technology)`

We **require** you to complete the below function. This function must take in a `technology` and return the **average** (i.e., **mean**) `net_winter_capacity` of **all** the power generators which use that particular `technology`. If a particular `technology` has **some** generators with **missing data**, those generators **must** be **ignored** while finding the average (i.e., ignored in **both** the *numerator* and the *denominator*).

This function can be **case-sensitive**. You **only** need to consider the technologies whose names **exactly match** `technology`.

```python
def avg_winter_capacity_of(technology):
    """
    avg_winter_capacity_of(technology)
    returns the average (mean) of `net_winter_capacity`
    of all power generators which use the given
    `technology`; generators with missing data
    are ignored.
    """    
    avg = 0
    counter = 0
    var = 0
    for idx in range(len(csv_rows)):
        if cell(idx, 'technology') == technology:
            if type(cell(idx, 'net_winter_capacity')) == float:
                var += cell(idx, 'net_winter_capacity')
                counter += 1
    if counter == 0:
        counter = 1
    avg = var / counter
    return avg
```

```python
grader.check("avg_winter_capacity_of")
```

Output:
```
avg_winter_capacity_of results: All test cases passed!
avg_winter_capacity_of - 1 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'not all rows are looped through (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
avg_winter_capacity_of - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'missing data is not ignored (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
avg_winter_capacity_of - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'function logic is incorrect (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
avg_winter_capacity_of - 4 message: Note that the Gradescope autograder will deduct points if your code 
```

**Question 18:** What is the **average** `net_winter_capacity` of all power generators that use the `technology` *Conventional Hydroelectric*?

You **must** use the `avg_winter_capacity_of` function to answer this question.

```python
# compute and store the answer in the variable 'hydro_avg_winter_capacity', then display it
hydro_avg_winter_capacity = avg_winter_capacity_of('Conventional Hydroelectric')
hydro_avg_winter_capacity
```

Output:
```
1.7491304347826093
```

```python
grader.check("q18")
```

Output:
```
q18 results: All test cases passed!
q18 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`avg_winter_capacity_of` function is not used to answer (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

**Question 19:** Which `technology` has the **highest** `net_winter_capacity` on **average**?

You **must** use the `avg_winter_capacity_of` function to answer this question. You do **not** have to worry about any ties. There is a **unique** technology with the highest `net_winter_capacity` on average.

**Hint**: You already created a list of the **unique** technologies to answer Question 1. Loop through that list and find which of those technologies has the **highest** average `net_winter_capacity` (using `avg_winter_capacity_of`).

```python

```

```python
# compute and store the answer in the variable 'max_winter_capacity_tech', then display it
technologies = []
list_with_duplicates = []
for idx in range(num_rows):
    list_with_duplicates.append(cell(idx, 'technology'))
for value in list_with_duplicates:
    if value not in technologies:
        technologies.append(value)
list_without_duplicates = []
for value in technologies:
    list_without_duplicates.append(value)
str1 = ''
technologies = list_without_duplicates
max_winter_capacity_tech = 0
highest_avg = 0
for val in technologies:
    if avg_winter_capacity_of(val) > highest_avg:
        highest_avg = avg_winter_capacity_of(val)
        str1 = val
max_winter_capacity_tech = str1
max_winter_capacity_tech
```

Output:
```
'Nuclear'
```

```python
grader.check("q19")
```

Output:
```
q19 results: All test cases passed!
q19 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'all technologies are looped through instead of just the unique technologies (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q19 - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'list of unique technologies is recomputed (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q19 - 4 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q19 - 5 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric poin
```

**Question 20:** Find the **difference** between the **average** net winter capacities of the `technology` with the **highest** average and the **second highest** average.

You must use the `avg_winter_capacity_of` function to answer this question.

**Hint**: You have already found the `technology` with the **highest** average. You need to find the `technology` with the **second highest** average, and find the difference between their average net winter capacities using the `avg_winter_capacity_of` function.

**Extra Hint**: There is also another (easier) way to solve this question. Try to find it! Reviewing Lab-P6 could help you find it.

```python
# compute and store the answer in the variable 'diff_avg_winter_capacity', then display it
diff_avg_winter_capacity = 0
technologies = []
list_with_duplicates = []
second_avg = 0
for idx in range(num_rows):
    list_with_duplicates.append(cell(idx, 'technology'))
for value in list_with_duplicates:
    if value not in technologies:
        technologies.append(value)
list_without_duplicates = []
for value in technologies:
    list_without_duplicates.append(value)
str1 = ''
technologies = list_without_duplicates
max_winter_capacity_tech = 0
highest_avg = 0
for val in technologies:
    if val == 'Nuclear':
        highest_avg = avg_winter_capacity_of('Nuclear')
    elif avg_winter_capacity_of(val) > second_avg:
        second_avg = avg_winter_capacity_of(val)
diff_avg_winter_capacity = highest_avg - second_avg
diff_avg_winter_capacity
```

Output:
```
271.00000000000006
```

```python
grader.check("q20")
```

Output:
```
q20 results: All test cases passed!
q20 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'all technologies are looped through instead of just the unique technologies (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q20 - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q20 - 4 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`avg_winter_capacity_of` function is not used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

```python
grader.check("general_deductions")
```

Output:
```
general_deductions results: All test cases passed!
general_deductions - 1 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'Did not save the notebook file prior to running the cell containing "export". We cannot see your output if you do not save before generating the zip file. This deduction will become stricter for future projects. (-3)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
general_deductions - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'Used concepts/modules such as dictionaries, or pandas not covered in class yet - built-in functions that you have been introduced to can be used. (-3)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
general_deductions - 3 message: Note that the Gradescope autograder will deduct points if your co
```

```python
grader.check("summary")
```

Output:
```
summary results: All test cases passed!
```

## Submission
It is recommended that at this stage, you Restart and Run all Cells in your notebook.
That will automatically save your work and generate a zip file for you to submit.

**SUBMISSION INSTRUCTIONS**:
1. **Upload** the zipfile to Gradescope.
2. If you completed the project with a **partner**, make sure to **add their name** by clicking "Add Group Member"
in Gradescope when uploading the zip file.
3. Check **Gradescope** results as soon as the auto-grader execution gets completed.
4. Your **final score** for this project is the score that you see on **Gradescope**.
5. You are **allowed** to resubmit on Gradescope as many times as you want to.
6. **Contact** a TA/PM if you lose any points on Gradescope for any **unclear reasons**.

```python
# running this cell will create a new save checkpoint for your notebook
from IPython.display import display, Javascript
display(Javascript('IPython.notebook.save_checkpoint();'))
```

```python
!jupytext --to py p6.ipynb
```

```python
public_tests.check_file_size("p6.ipynb")
grader.export(pdf=False, run_tests=False, files=["p6.py"])
```

Tests for avg_winter_capacity_of (points: 0)
```
>>> 
>>> public_tests.rubric_check('avg_winter_capacity_of: not all rows are looped through')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'not all rows are looped through (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('avg_winter_capacity_of: missing data is not ignored')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'missing data is not ignored (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('avg_winter_capacity_of: function logic is incorrect')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'function logic is incorrect (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('avg_winter_capacity_of: `cell` function is not used to read data', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`cell` function is not used to read data (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for cell (points: 0)
```
>>> 
>>> public_tests.rubric_check('cell: variables `csv_data`, `csv_header`, and `csv_rows` are not defined as expected')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'variables `csv_data`, `csv_header`, and `csv_rows` are not defined as expected (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('cell: function does not return `None` when data is missing')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'function does not return `None` when data is missing (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('cell: function does not typecast to `int` or `float` when expected')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'function does not typecast to `int` or `float` when expected (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('cell: function logic is incorrect', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'function logic is incorrect (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('cell: `process_csv` function is not used to read `power_generators.csv`')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`process_csv` function is not used to read `power_generators.csv` (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for find_entities_with_phrase (points: 0)
```
>>> 
>>> public_tests.rubric_check('find_entities_with_phrase: not all rows are looped through')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'not all rows are looped through (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('find_entities_with_phrase: case insensitive matches are not considered')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'case insensitive matches are not considered (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('find_entities_with_phrase: duplicates are not removed')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'duplicates are not removed (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('find_entities_with_phrase: `cell` function is not used to read data', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`cell` function is not used to read data (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for find_indices_within (points: 0)
```
>>> 
>>> public_tests.rubric_check('find_indices_within: not all rows are looped through')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'not all rows are looped through (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('find_indices_within: incorrect comparison operators are used')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect comparison operators are used (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('find_indices_within: function logic is incorrect')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'function logic is incorrect (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('find_indices_within: `cell` function is not used to read data', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`cell` function is not used to read data (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for general_deductions (points: 0)
```
>>> 
>>> public_tests.rubric_check('general_deductions: Did not save the notebook file prior to running the cell containing "export". We cannot see your output if you do not save before generating the zip file. This deduction will become stricter for future projects.')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'Did not save the notebook file prior to running the cell containing "export". We cannot see your output if you do not save before generating the zip file. This deduction will become stricter for future projects. (-3)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('general_deductions: Used concepts/modules such as dictionaries, or pandas not covered in class yet - built-in functions that you have been introduced to can be used.')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'Used concepts/modules such as dictionaries, or pandas not covered in class yet - built-in functions that you have been introduced to can be used. (-3)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('general_deductions: Import statements are not all placed at the top of the notebook.')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'Import statements are not all placed at the top of the notebook. (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for median (points: 0)
```
>>> 
>>> public_tests.rubric_check('median: function modifies the original input list')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'function modifies the original input list (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('median: function logic is incorrect when inputs are of odd length')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'function logic is incorrect when inputs are of odd length (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('median: function logic is incorrect when inputs are of even length')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'function logic is incorrect when inputs are of even length (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for num_generators_by (points: 0)
```
>>> 
>>> public_tests.rubric_check('num_generators_by: not all rows are looped through')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'not all rows are looped through (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('num_generators_by: function logic is incorrect')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'function logic is incorrect (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('num_generators_by: `cell` function is not used to read data', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`cell` function is not used to read data (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for q1 (points: 0)
```
>>> public_tests.check('q1', technologies)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q1: not all rows are looped through')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'not all rows are looped through (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q1: incorrect logic is used')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q1: `cell` function is not used to read data', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`cell` function is not used to read data (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> public_tests.rubric_check('q1: public tests')
All test cases passed!
```


Tests for q10 (points: 0)
```
>>> public_tests.check('q10', num_generators_by_river)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q10: incorrect logic is used to answer')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q10: `num_generators_by` function is not used to answer', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`num_generators_by` function is not used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q10: `find_entities_with_phrase` function is not used to answer', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`find_entities_with_phrase` function is not used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for q11 (points: 0)
```
>>> public_tests.check('q11', most_generators_entity)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q11: all entities are looped through instead of just the unique entities')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'all entities are looped through instead of just the unique entities (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q11: incorrect logic is used to answer')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q11: `num_generators_by` function is not used to answer', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`num_generators_by` function is not used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> public_tests.rubric_check('q11: public tests')
All test cases passed!
```


Tests for q12 (points: 0)
```
>>> public_tests.check('q12', num_generators_in_milwaukee)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q12: length of the list is found by looping through it')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'length of the list is found by looping through it (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q12: `find_indices_within` function is not used to answer', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`find_indices_within` function is not used to answer (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for q13 (points: 0)
```
>>> public_tests.check('q13', uw_madison_technologies)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q13: duplicates are not removed from the list')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'duplicates are not removed from the list (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q13: incorrect logic is used to answer')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q13: `find_indices_within` function is not used to answer', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`find_indices_within` function is not used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for q14 (points: 0)
```
>>> public_tests.check('q14', north_wisconsin_most_powerful)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q14: incorrect logic is used to answer')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q14: `find_indices_within` function is not used to answer', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`find_indices_within` function is not used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q14: `cell` function is not used to read data', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`cell` function is not used to read data (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> public_tests.rubric_check('q14: public tests')
All test cases passed!
```


Tests for q15 (points: 0)
```
>>> public_tests.check('q15', winnebago_hydro_winter_capacity)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q15: incorrect logic is used to answer')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q15: `find_indices_within` function is not used to answer', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`find_indices_within` function is not used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q15: `median` function is not used to answer', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`median` function is not used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q15: `cell` function is not used to read data', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`cell` function is not used to read data (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for q16 (points: 0)
```
>>> public_tests.check('q16', point_beach_summer_capacity)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q16: `total_summer_capacity_of` function is not used to answer')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`total_summer_capacity_of` function is not used to answer (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for q17 (points: 0)
```
>>> public_tests.check('q17', median_summer_capacity)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q17: all power plants are looped through instead of just the unique power plants')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'all power plants are looped through instead of just the unique power plants (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q17: incorrect logic is used to answer')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q17: `total_summer_capacity_of` function is not used to answer', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`total_summer_capacity_of` function is not used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q17: `median` function is not used to answer', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`median` function is not used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> public_tests.rubric_check('q17: public tests')
All test cases passed!
```


Tests for q18 (points: 0)
```
>>> public_tests.check('q18', hydro_avg_winter_capacity)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q18: `avg_winter_capacity_of` function is not used to answer')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`avg_winter_capacity_of` function is not used to answer (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for q19 (points: 0)
```
>>> public_tests.check('q19', max_winter_capacity_tech)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q19: all technologies are looped through instead of just the unique technologies')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'all technologies are looped through instead of just the unique technologies (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q19: list of unique technologies is recomputed')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'list of unique technologies is recomputed (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q19: incorrect logic is used to answer')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q19: `avg_winter_capacity_of` function is not used to answer', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`avg_winter_capacity_of` function is not used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> public_tests.rubric_check('q19: public tests')
All test cases passed!
```


Tests for q2 (points: 0)
```
>>> public_tests.check('q2', count_dane)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q2: not all rows are looped through')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'not all rows are looped through (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q2: incorrect logic is used')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q2: `cell` function is not used to read data', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`cell` function is not used to read data (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for q20 (points: 0)
```
>>> public_tests.check('q20', diff_avg_winter_capacity)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q20: all technologies are looped through instead of just the unique technologies')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'all technologies are looped through instead of just the unique technologies (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q20: incorrect logic is used to answer')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q20: `avg_winter_capacity_of` function is not used to answer', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`avg_winter_capacity_of` function is not used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> public_tests.rubric_check('q20: public tests')
All test cases passed!
```


Tests for q3 (points: 0)
```
>>> public_tests.check('q3', total_summer_capacity)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q3: not all rows are looped through')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'not all rows are looped through (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q3: incorrect logic is used')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q3: `cell` function is not used to read data', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`cell` function is not used to read data (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for q4 (points: 0)
```
>>> public_tests.check('q4', madison_entities)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q4: `find_entities_with_phrase` function is not used to answer')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`find_entities_with_phrase` function is not used to answer (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for q5 (points: 0)
```
>>> public_tests.check('q5', entities_contain_wisconsin_power)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q5: incorrect logic is used to answer')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q5: `cell` function is not used to read data', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`cell` function is not used to read data (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> public_tests.rubric_check('q5: public tests')
All test cases passed!
```


Tests for q6 (points: 0)
```
>>> public_tests.check('q6', entities_contain_solar_llc)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q6: incorrect logic is used to answer')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q6: `cell` function is not used to read data', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`cell` function is not used to read data (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> public_tests.rubric_check('q6: public tests')
All test cases passed!
```


Tests for q7 (points: 0)
```
>>> public_tests.check('q7', plant_50614_generators)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q7: not all rows are looped through')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'not all rows are looped through (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q7: the output is not sorted correctly')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'the output is not sorted correctly (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q7: incorrect logic is used', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q7: `cell` function is not used to read data', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`cell` function is not used to read data (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> public_tests.rubric_check('q7: public tests')
All test cases passed!
```


Tests for q8 (points: 0)
```
>>> public_tests.check('q8', powerful_hydro_electric_plants)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q8: not all rows are looped through')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'not all rows are looped through (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q8: the output is not sorted correctly')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'the output is not sorted correctly (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q8: incorrect comparison operators are used')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect comparison operators are used (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q8: incorrect logic is used', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q8: `cell` function is not used to read data', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`cell` function is not used to read data (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for q9 (points: 0)
```
>>> public_tests.check('q9', num_generators_by_mge)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q9: `num_generators_by` function is not used to answer')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`num_generators_by` function is not used to answer (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for summary (points: 127)
```
>>> public_tests.get_summary()
Total Score: 100/100
```


Tests for total_summer_capacity_of (points: 0)
```
>>> 
>>> public_tests.rubric_check('total_summer_capacity_of: not all rows are looped through')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'not all rows are looped through (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('total_summer_capacity_of: missing data is not ignored')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'missing data is not ignored (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('total_summer_capacity_of: function logic is incorrect')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'function logic is incorrect (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('total_summer_capacity_of: `cell` function is not used to read data', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`cell` function is not used to read data (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
