```python
# import and initialize otter
import otter
grader = otter.Notebook("p8.ipynb")
```

```python
import public_tests
```

```python
# PLEASE FILL IN THE DETAILS
# enter none if you don't have a project partner
# you will have to add your partner as a group member on Gradescope even after you fill this

# project: p8
# submitter: alghazzy
# partner: jmbruggink
# hours: 4
```

# Project 8: Going to the Movies

## Learning Objectives:

In this project, you will demonstrate how to:

* integrate relevant information from various sources (e.g. multiple csv files),
* build appropriate data structures for organized and informative presentation (e.g. list of dictionaries),
* practice good coding style.

Please go through [Lab-P8](https://git.doit.wisc.edu/cdis/cs/courses/cs220/cs220-f23-projects/-/tree/main/lab-p8) before working on this project. The lab introduces some useful techniques related to this project.

## Testing your code:

Along with this notebook, you must have downloaded the file `public_tests.py`. If you are curious about how we test your code, you can explore this file, and specifically the function `get_expected_json`, to understand the expected answers to the questions.

<h2 style="color:red">Warning (Note on Academic Misconduct):</h2>

**IMPORTANT**: **P8 and P9 are two parts of the same data analysis.** You **cannot** switch project partners between these two projects. That is if you partner up with someone for P8, you have to sustain that partnership until the end of P9. Now may be a good time to review [our course policies](https://cs220.cs.wisc.edu/f23/syllabus.html).

Under any circumstances, **no more than two students are allowed to work together on a project** as mentioned in the course policies. If your code is flagged by our code similarity detection tools, **both partners will be responsible** for sharing/copying the code, even if the code is shared/copied by one of the partners with/from other non-partner student(s). Note that each case of plagiarism will be reported to the Dean of Students with a zero grade on the project. **If you think that someone cannot be your project partner then don’t make that student your lab partner.**

**<font color = "red">Project partners must submit only one copy of their project on Gradescope, but they must include the names of both partners.</font>**

## Introduction:

In this project and the next, we will be working on the [IMDb Movies Dataset](https://www.imdb.com/interfaces/). We will use Python to discover some cool facts about our favorite movies, cast, and directors.

In this project, you will combine the data from the movie and mapping files into a more useful format.
Start by downloading the following files: `public_tests.py`, `small_mapping.csv`, `small_movies.csv`, `mapping.csv`, and `movies.csv`.

## The Data:

Open `movies.csv` and `mapping.csv` in any spreadsheet viewer, and see what the data looks like.
`movies.csv` has ~75,000 rows and `mapping.csv` has ~220,000 rows. These files store information about **every** movie on the IMDb dataset which was released in the US, and received over **250** reviews on IMDb. These datasets are **very** large when compared to `small_movies.csv` and `small_mapping.csv` from [Lab-P8](https://git.doit.wisc.edu/cdis/cs/courses/cs220/cs220-f23-projects/-/tree/main/lab-p8), but the data is stored in the **same format**. For a description of the datasets, please refer back to [Lab-P8](https://git.doit.wisc.edu/cdis/cs/courses/cs220/cs220-f23-projects/-/tree/main/lab-p8).

Before we start working with these very large datasets, let us start with the much smaller datasets, `small_movies.csv` and `small_mapping.csv` from Lab-P8. In the latter half of P8 and in P9, you will be working with `movies.csv` and `mapping.csv`. Since the files `movies.csv` and `mapping.csv` are large, some of the functions you write in P8 and P9 **may take a while to execute**. You do not have to panic if a single cell takes between 5 to 10 seconds to run. If any cell takes significantly longer, follow the recommendations below:

- **Do not** call **slow functions** multiple times within a loop.
- **Do not** call functions that **iterate over the entire dataset within a loop**; instead, call the function before the loop and store the result in a variable.
- **Do not** compute quantities **inside a loop** if it can be computed outside the loop; for example, if you want to calculate the average of a list, you should use the loop to find the numerator and denominator but divide **once** after the loop ends instead of inside the loop.

## Project Requirements:

You **may not** hardcode indices in your code, unless the question explicitly asks you to do so. If you open your `.csv` files with Excel, manually count through the rows and use this number to loop through the dataset, this is also considered as hardcoding. If any instances of hardcoding are found during code review, the Gradescope autograder will **deduct** points from your public score.

**Store** your final answer for each question in the **variable specified for each question**. This step is important because Otter grades your work by comparing the value of this variable against the correct answer.

For some of the questions, we'll ask you to write (then use) a function to compute the answer. If you compute the answer **without** creating the function we ask you to write, the Gradescope autograder will **deduct** points from your public score, even if the way you did it produced the correct answer.

#### Required Functions:
- `get_mapping`
- `get_raw_movies`
- `get_movies`
- `find_specific_movies`
- `bucketize_by_genre`

In this project, you will also be required to define certain **data structures**. If you do not create these data structures exactly as specified, the Gradescope autograder will **deduct** points from your public score, even if the way you did it produced the correct answer.

#### Required Data Structures:
- `small_movies_data`
- `movies`
- `genre_dict`

You are only allowed to define these data structures **once** and the Gradescope autograder will **deduct** points from your public score if you redefine the values of these variables.

In this project (and the next), you will be asked to create **lists** of movies. For all such questions, **unless it is explicitly mentioned otherwise**, the movies should be in the **same order** as in the `movies.csv` (or `small_movies.csv`) file. Similarly, for each movie, the **list** of `genres`, `directors`, and `cast` members should always be in the **same order** as in the `movies.csv` (or `small_movies.csv`) file.

Students are only allowed to use Python commands and concepts that have been taught in the course prior to the release of P8. Therefore, you should not use the pandas module. The Gradescope autograder will **deduct** points from your public score otherwise.

In addition, you are also **required** to follow the requirements below:
- **Do not use the method `csv.DictReader` for P8**. Although the required output can be obtained using this method, one of the learning outcomes of this project is to demonstrate your ability to build dictionaries with your own code.  
- Additional import statements beyond those that are stated in the directions are not allowed. For this project, we allow you to use `csv` and `copy` packages (that is, you can use the `import csv` and `import copy` statements in your submission). You should not use concepts / modules that are yet to be covered in this course; for example: you should **not** use modules like `pandas`.

The Gradescope autograder will **deduct** points accordingly, if you don't follow the provided directions.

For more details on what will cause you to lose points during code review and specific requirements, please take a look at the [Grading rubric](https://git.doit.wisc.edu/cdis/cs/courses/cs220/cs220-f23-projects/-/blob/main/p8/rubric.md).

## Questions and Functions:

Let us start by importing all the modules we will need for this project.



```python
# it is considered a good coding practice to place all import statements at the top of the notebook
# please place all your import statements in this cell if you need to import any more modules for this project
import csv
import math
import copy
```

### Function 1: `get_mapping(path)`

We require you to complete the below function to answer the next several questions (this is a **requirement**, and you will **lose points** if you do not implement this function). You may copy/paste code from your lab-p8 notebook to finish this function.

```python
def process_csv(filename):
    example_file = open(filename, encoding="utf-8")
    example_reader = csv.reader(example_file)
    example_data = list(example_reader)
    example_file.close()
    return example_data
def get_mapping(path):
    csv_data = process_csv(path)
    csv_header = csv_data[0]
    rows = csv_data[0:]
    dict_x = {}
    
    for row in rows:
        dict_x[row[0]] = row[1]
        
    return dict_x
```

```python
grader.check("get_mapping")
```

Output:
```
get_mapping results: All test cases passed!
get_mapping - 1 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'function logic is incorrect (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
get_mapping - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'function is defined more than once (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

**Question 1:** What is returned by `get_mapping("small_mapping.csv")`?

Your output **must** be a **dictionary** which maps the *IDs* in `small_mapping.csv` to *names*.

```python
# compute and store the answer in the variable 'small_mapping', then display it
small_mapping = get_mapping("small_mapping.csv")
small_mapping
```

Output:
```
{'tt15398776': 'Oppenheimer',
 'tt1517268': 'Barbie',
 'nm0634240': 'Christopher Nolan',
 'nm0614165': 'Cillian Murphy',
 'nm1289434': 'Emily Blunt',
 'nm0000354': 'Matt Damon',
 'nm0000375': 'Robert Downey Jr.',
 'nm1950086': 'Greta Gerwig',
 'nm3053338': 'Margot Robbie',
 'nm0331516': 'Ryan Gosling',
 'nm4793987': 'Issa Rae',
 'nm0571952': 'Kate McKinnon'}
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
q1 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`get_mapping` function is not used to answer (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

**Question 2:** What is the **value** associated with the **key** *nm1289434*?

Your output **must** be a **string**. You **must** use the variable `small_mapping` defined above to answer this question. You **must** not call the function `get_mapping` again on this dataset.

```python
# access and store the answer in the variable 'nm1289434_value', then display it
nm1289434_value = small_mapping['nm1289434']
nm1289434_value
```

Output:
```
'Emily Blunt'
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
q2 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`small_mapping` data structure is not used to answer (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q2 - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'answer uses loops to iterate over dictionary keys (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

**Question 3:** What are the **values** associated with **keys** that **begin** with *nm*?

Your output **must** be a **list** of **strings**. You **must** find **only** the values of the keys that **begin** with *nm*, and **not** the keys that contain *nm*. You **must** use the variable `small_mapping` defined above to answer this question.

```python
# compute and store the answer in the variable 'nm_values', then display it
nm_values = []
for value in small_mapping:
    if value[:2] == 'nm':
        nm_values.append(small_mapping[value])
nm_values
```

Output:
```
['Christopher Nolan',
 'Cillian Murphy',
 'Emily Blunt',
 'Matt Damon',
 'Robert Downey Jr.',
 'Greta Gerwig',
 'Margot Robbie',
 'Ryan Gosling',
 'Issa Rae',
 'Kate McKinnon']
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
q3 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'answer does not check only for keys that begin with *nm* (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q3 - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`small_mapping` data structure is not used to answer (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

**Question 4:** Find the **keys** of the people (keys **beginning** with *nm*) whose **last name** is *Murphy*.

Your output **must** be a **list** of **string(s)**. You **must** use the variable `small_mapping` defined above to answer this question.

**Requirements:** Your **code** must be robust and satisfy all the requirements, even if you were to run this on a larger dataset (such as `mapping.csv`). In particular:
1. You will **lose points** if your code would find people whose **first** name or **middle** name is *Murphy* (e.g. *Murphy Patrick Martin*, *Lisa Murphy Collins*).
2. You will **lose points** if your code would find people whose **last** name contains *Murphy* as a **substring** (e.g. *Benjamin Harter-Murphy*). The name should be **exactly** *Murphy*. 
3. You will **lose points** if your code would find any **movie titles** (e.g. *Sergeant Murphy*).

```python
# compute and store the answer in the variable 'nm_murphy', then display it
nm_murphy = []
for key, value in small_mapping.items():
    x = value.split()
    if key[:2] == 'nm':
        if len(x) > 0 and x[-1] == 'Murphy':
            nm_murphy.append(key)
nm_murphy
```

Output:
```
['nm0614165']
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
q4 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'answer does not exclude people whose first name or middle name is Murphy (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q4 - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'answer does not exclude people whose last name contains Murphy as a substring (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q4 - 4 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'answer does not exclude movie titles (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q4 - 5 message: Note that the Gradescope autograder will deduct points if your code 
```

#### Now, let's move on to reading the movie files!

### Function 2: `get_raw_movies(path)`

We require you to complete the below function to answer the next several questions (this is a **requirement**, and you will **lose points** if you do not implement this function).

This function **must** return a **list** of **dictionaries**, where each **dictionary** is of the following format:

```python
   {
        'title': <title-id>,
        'year': <the year as an integer>,
        'duration': <the duration as an integer>,
        'genres': [<genre1>, <genre2>, ...],
        'rating': <the rating as a float>,
        'directors': [<director-id1>, <director-id2>, ...],
        'cast': [<actor-id1>, <actor-id2>, ....]
    }
```

Here is an example:

```python
    {
        'title': 'tt0033313',
        'year': 1941,
        'duration': 59,
        'genres': ['Western'],
        'rating': 5.2,
        'directors': ['nm0496505'],
        'cast': ['nm0193318', 'nm0254381', 'nm0279961', 'nm0910294', 'nm0852305']
    }
```

You may copy/paste code from your lab-p8 notebook to finish this function.

```python
def get_raw_movies(path):
    raw_movies_list = []
    with open(path, mode='r', newline='', encoding='utf-8') as file:
        reader = csv.reader(file)
        csv_header = next(reader)
        for row in reader:
            movie = {'title': row[csv_header.index('title')], 'year': int(row[csv_header.index('year')]), 'duration': int(row[csv_header.index('duration')]), 'genres': row[csv_header.index('genres')].split(', '), 'rating': float(row[csv_header.index('rating')]), 'directors': row[csv_header.index('directors')].split(', '), 'cast': row[csv_header.index('cast')].split(', ')}
                    
            raw_movies_list.append(movie)
    return raw_movies_list
    
```

```python
grader.check("get_raw_movies")
```

Output:
```
get_raw_movies results: All test cases passed!
get_raw_movies - 1 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'function logic is incorrect (-3)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
get_raw_movies - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'column indices are hardcoded instead of using the header to identify column names (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
get_raw_movies - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'function is defined more than once (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

**Question 5:** What is returned by `get_raw_movies("small_movies.csv")`?

Your output **must** be a **list** of **dictionaries** where each dictionary contains information about a movie.

```python
# compute and store the answer in the variable 'raw_small_movies', then display it
raw_small_movies = get_raw_movies("small_movies.csv")
raw_small_movies
```

Output:
```
[{'title': 'tt15398776',
  'year': 2023,
  'duration': 180,
  'genres': ['Biography', 'Drama', 'History'],
  'rating': 8.7,
  'directors': ['nm0634240'],
  'cast': ['nm0614165', 'nm1289434', 'nm0000354', 'nm0000375']},
 {'title': 'tt1517268',
  'year': 2023,
  'duration': 114,
  'genres': ['Adventure', 'Comedy', 'Fantasy'],
  'rating': 7.5,
  'directors': ['nm1950086'],
  'cast': ['nm3053338', 'nm0331516', 'nm4793987', 'nm0571952']}]
```

```python
grader.check("q5")
```

Output:
```
q5 results: All test cases passed!
q5 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`get_raw_movies` function is not used to answer (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

If your answer looks correct, but does not pass `grader.check`, make sure that the **datatypes** are all correct. Also make sure that the **directors** and **cast**  are in the **same order** as in `small_movies.csv`.

**Question 6:** How **many** cast members does the **first** movie have?

Your output **must** be an **int**. You **must** use the variable `raw_small_movies` defined above to answer this question. You **must** not call the function `get_raw_movies` again on this dataset.

```python
num_cast_first_movie = len(raw_small_movies[0]['cast'])
num_cast_first_movie
```

Output:
```
4
```

```python
grader.check("q6")
```

Output:
```
q6 results: All test cases passed!
q6 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'answer uses loops to find number of cast members (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q6 - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'answer uses loops to iterate over dictionary keys (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q6 - 4 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q6 - 5 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`raw_small_movies` 
```

**Question 7:** What is the *ID* of the **first** cast member listed for the **first** movie of the dataset?

Your output **must** be a **string**. You **must** use the variable `raw_small_movies` defined above to answer this question.

```python
# compute and store the answer in the variable 'first_actor_id_first_movie', then display it
first_actor_id_first_movie = raw_small_movies[0]['cast'][0]
first_actor_id_first_movie
```

Output:
```
'nm0614165'
```

```python
grader.check("q7")
```

Output:
```
q7 results: All test cases passed!
q7 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q7 - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'answer uses loops to iterate over dictionary keys (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q7 - 4 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`raw_small_movies` data structure is not used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

### Function 3: `get_movies(movies_path, mapping_path)`

We require you to complete the below function to answer the next several questions (this is a **requirement**, and you will **lose points** if you do not implement this function).


This function **must** return a **list** of **dictionaries**, where each **dictionary** is of the following format:

```python
   {
        'title': "the movie name",
        'year': <the year as an integer>,
        'duration': <the duration as an integer>,
        'genres': [<genre1>, <genre2>, ...],
        'rating': <the rating as a float>,
        'directors': ["director-name1", "director-name2", ...],
        'cast': ["actor-name1", "actor-name2", ....]
    }
```

Here is an example:

```python
    {
        'title': 'Across the Sierras',
        'year': 1941,
        'duration': 59,
        'genres': ['Western'],
        'rating': 5.2,
        'directors': ['D. Ross Lederman'],
        'cast': ['Dick Curtis', 'Bill Elliott', 'Richard Fiske', 'Luana Walters', 'Dub Taylor']
    }
```

You may copy/paste code from your Lab-P8 notebook to finish this function.

```python
def get_movies(movies_path, mapping_path):
    map_dict = get_mapping(mapping_path)
    raw_movies = get_raw_movies(movies_path)
    movies = []
    for raw_movie in raw_movies:
        movie = {} 
        movie['title'] = map_dict[raw_movie['title']]
        movie['year'] = int(raw_movie['year'])
        movie['duration'] = int(raw_movie['duration'])
        movie['genres'] = raw_movie['genres']
        movie['rating'] = float(raw_movie['rating'])
        movie['directors'] = [map_dict[director] for director in raw_movie['directors']]
        movie['cast'] = [map_dict[actor] for actor in raw_movie['cast']]
        movies.append(movie)

    return movies
```

```python
grader.check("get_movies")
```

Output:
```
get_movies results: All test cases passed!
get_movies - 1 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'function is called more than once with the same dataset (-4)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
get_movies - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'function logic is incorrect (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
get_movies - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`get_mapping` and `get_raw_movies` functions are not used (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
get_movies - 4 message: Note that the Gradescope autograder will deduct points if your code fails 
```

**Question 8:** What is returned by `get_movies("small_movies.csv", "small_mapping.csv")`?

Your output **must** be a **list** of **dictionaries** where each dictionary contains information about a movie.

```python
# compute and store the answer in the variable 'small_movies', then display it
small_movies = get_movies("small_movies.csv", "small_mapping.csv")
small_movies
```

Output:
```
[{'title': 'Oppenheimer',
  'year': 2023,
  'duration': 180,
  'genres': ['Biography', 'Drama', 'History'],
  'rating': 8.7,
  'directors': ['Christopher Nolan'],
  'cast': ['Cillian Murphy',
   'Emily Blunt',
   'Matt Damon',
   'Robert Downey Jr.']},
 {'title': 'Barbie',
  'year': 2023,
  'duration': 114,
  'genres': ['Adventure', 'Comedy', 'Fantasy'],
  'rating': 7.5,
  'directors': ['Greta Gerwig'],
  'cast': ['Margot Robbie', 'Ryan Gosling', 'Issa Rae', 'Kate McKinnon']}]
```

```python
grader.check("q8")
```

Output:
```
q8 results: All test cases passed!
q8 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`get_movies` function is not used to answer (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

```python
grader.check("small_movies")
```

Output:
```
small_movies results: All test cases passed!
small_movies - 1 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'data structure is defined incorrectly (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
small_movies - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'data structure is redefined or modified after its initialization (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

**Question 9:** What is `title` of the **second** movie in `small_movies`?

Your output **must** be a **string**. You **must** use the variable `small_movies` defined above to answer this question. You **must not** call the function `get_movies` again on this dataset.

```python
# compute and store the answer in the variable 'second_movie_title_small_movies', then display it
second_movie_title_small_movies = small_movies[1]['title']
second_movie_title_small_movies
```

Output:
```
'Barbie'
```

```python
grader.check("q9")
```

Output:
```
q9 results: All test cases passed!
q9 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q9 - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'answer uses loops to iterate over dictionary keys (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q9 - 4 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`small_movies` data structure is not used to answer (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

**Question 10:** Who are the `cast` members of the **second** movie in `small_movies`?

Your output **must** be a **list** of **string(s)**. You **must** use the variable `small_movies` defined above to answer this question.

```python
# compute and store the answer in the variable 'second_movie_cast_small_movies', then display it
second_movie_cast_small_movies = small_movies[1]['cast']
second_movie_cast_small_movies
```

Output:
```
['Margot Robbie', 'Ryan Gosling', 'Issa Rae', 'Kate McKinnon']
```

```python
grader.check("q10")
```

Output:
```
q10 results: All test cases passed!
q10 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q10 - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'answer uses loops to iterate over dictionary keys (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q10 - 4 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`small_movies` data structure is not used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

**Question 11:** Who are the `directors` of the **last** movie in `small_movies`?

Your output **must** be a **list** of **string(s)**. You **must** use the variable `small_movies` defined above to answer this question.

```python
# compute and store the answer in the variable 'last_movie_directors_small_movies', then display it
last_movie_directors_small_movies = small_movies[-1]['directors']
last_movie_directors_small_movies
```

Output:
```
['Greta Gerwig']
```

```python
grader.check("q11")
```

Output:
```
q11 results: All test cases passed!
q11 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q11 - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'answer uses loops to iterate over dictionary keys (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q11 - 4 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`small_movies` data structure is not used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

#### Now that you’ve made it this far, your functions must be working pretty well with small datasets. Next, let's try a much bigger dataset!

Run the following code to open the full dataset:

```python
movies = get_movies("movies.csv", "mapping.csv")
len(movies)
```

Output:
```
73075
```

As the files are very large, this cell is expected to take around five seconds to run. If it takes much longer (say, around thirty seconds), then you will **need** to **optimize** your `get_movies` function so it runs faster.

**Warning**: You are **not** allowed to call `get_movies` more than once on the full dataset (`movies.csv` and `mapping.csv`) in your notebook. Instead, reuse the `movies` variable, which is more efficient. The Gradescope autograder will deduct **points** if you call `get_movies` again on these files.

**Warning:** Do **not** display the value of the variable `movies` **anywhere** in your notebook. It will take up a **lot** of space, and your **Gradescope code will not be displayed** for grading. Instead you should verify `movies` has the correct value by looking at a small *slice* of the **list** as in the question below. 

```python
grader.check("movies")
```

Output:
```
movies results: All test cases passed!
movies - 1 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'data structure is defined incorrectly (-3)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
movies - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'data structure is redefined or modified after its initialization (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

**Question 12:** What are the movies in `movies[22202:22220]`?

Your answer should be a *list* of *dictionaries* that follows the format below:

```python
[{'title': 'Bad Night',
  'year': 2015,
  'duration': 92,
  'genres': ['Adventure', 'Comedy', 'Crime'],
  'rating': 4.9,
  'directors': ['Chris Riedell', 'Nick Riedell'],
  'cast': ['Lauren Elizabeth', 'Jenn McAllister', 'Julianna Guill', 'udy Marte']},
 {'title': 'Operation Camel',
  'year': 1960,
  'duration': 100,
  'genres': ['Comedy'],
  'rating': 5.7,
  'directors': ['Sven Methling'],
  'cast': ['Paul Hagen', 'Louis Miehe-Renard', 'Ebbe Langberg', 'Preben Kaas']},
 {'title': 'True Story',
  'year': 2015,
  'duration': 99,
  'genres': ['Biography', 'Crime', 'Drama'],
  'rating': 6.3,
  'directors': ['Rupert Goold'],
  'cast': ['James Franco', 'Jonah Hill', 'Felicity Jones', 'Maria Dizzia']},
 ...]
```

```python
# compute and store the answer in the variable 'movies_22202_22220', then display it
movies_22202_22220 = movies[22202:22220]
movies_22202_22220
```

Output:
```
[{'title': 'Bad Night',
  'year': 2015,
  'duration': 92,
  'genres': ['Adventure', 'Comedy', 'Crime'],
  'rating': 4.9,
  'directors': ['Chris Riedell', 'Nick Riedell'],
  'cast': ['Lauren Elizabeth',
   'Jenn McAllister',
   'Julianna Guill',
   'Judy Marte']},
 {'title': 'Operation Camel',
  'year': 1960,
  'duration': 100,
  'genres': ['Comedy'],
  'rating': 5.7,
  'directors': ['Sven Methling'],
  'cast': ['Paul Hagen',
   'Louis Miehe-Renard',
   'Ebbe Langberg',
   'Preben Kaas']},
 {'title': 'True Story',
  'year': 2015,
  'duration': 99,
  'genres': ['Biography', 'Crime', 'Drama'],
  'rating': 6.3,
  'directors': ['Rupert Goold'],
  'cast': ['James Franco', 'Jonah Hill', 'Felicity Jones', 'Maria Dizzia']},
 {'title': 'The Death of Louis XIV',
  'year': 2016,
  'duration': 115,
  'genres': ['Biography', 'Drama', 'History'],
  'rating': 6.8,
  'directors': ['Albert Serra'],
  'cast': ['Jean-Pierre Léaud',
   "Patrick d'Assumçao",
   'Marc Susini',
   'Bernard Belin']},
 {'title'
```

```python
grader.check("q12")
```

Output:
```
q12 results: All test cases passed!
q12 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`movies` data structure is not used to answer (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

**Question 13:** What is the **number** of movies released so far in the `year` *2023*?

Your outuput must be an **int**. You **must** use the variable `movies` defined above to answer this question.

```python
# compute and store the answer in the variable 'num_movies_2023', then display it
num_movies_2023 = 0
for movie in movies:
    if movie['year'] == 2023:
        num_movies_2023 += 1
num_movies_2023
```

Output:
```
1179
```

```python
grader.check("q13")
```

Output:
```
q13 results: All test cases passed!
q13 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q13 - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'answer uses loops to iterate over dictionary keys (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q13 - 4 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`movies` data structure is not used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

### Function 4: `find_specific_movies(movies, keyword)`

Now that we have created this data structure `movies`, we can start doing some fun things with the data!
We will continue working on this data structure for the next project (P9) as well.

Let us now use this data structure `movies` to create a **search bar** like the one in Netflix!
**Do not change the below function in any way**.
This function takes in a keyword like a substring of a title, a genre, or the name of a person, and returns a list of relevant movies with that title, genre, or cast member/director.

**Warning:** As `movies` is very large, the function `find_specific_movies` may take five to ten seconds to run. This is normal and you should not panic if it takes a while to run.

```python
# DO NOT EDIT THIS CELL AND DO NOT REDEFINE THIS FUNCTION ELSEWHERE
def find_specific_movies(movies, keyword):
    """
    find_specific_movies(movies, keyword) takes a list of movie dictionaries 
    and a keyword; it returns a list of movies that contain the keyword
    in either its title, genre, cast or directors.
    """
    idx = 0
    while idx < len(movies):
        movie = movies[idx]
        # note: \ enables you split a long line of code into two lines
        if (keyword not in movie['title']) and (keyword not in movie["genres"]) \
        and (keyword not in movie["directors"]) and (keyword not in movie["cast"]):
            movies.pop(idx)
        else:
            idx += 1
    return movies
```

**Important:** While it might look as if we are making it easy for you by providing `find_specific_movies`, there is a catch! There is a subtle flaw with the way the function is defined, that will cause you issues in the next two questions. If you can spot this flaw by just observing the definition of `find_specific_movies`, congratulations! Since you are **not** allowed to modify the function definition, you will have to be a little clever with your function arguments to sidestep the flaw with the function definition.

If you don't see anything wrong with the function just yet, don't worry about it. Solve Question 14 and Question 15 as you normally would, and see if you notice anything suspicious about your answers.

```python
grader.check("find_specific_movies")
```

Output:
```
find_specific_movies results: All test cases passed!
find_specific_movies - 1 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'function is redefined or defined more than once (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

**Question 14:** List all the movies involving *Alfred Hitchcock*.

Your answer **must** be a **list** of **dictionaries**.

You **must** answer this question by calling `find_specific_movies` with the keyword `"Alfred Hitchcock"`.

The `find_specific_movies` function is expected to take around 5 seconds or more to run, so do not panic if it takes so long to run.

Remember that you are **not** allowed to modify the definition of `find_specific_movies`. You will need to cleverly pass arguments to `find_specific_movies` (in both Question 14 and Question 15) to ensure that `movies` does not get modified by the function calls. Take a look at the Lecture Slides from October 25 for more hints. You will have to Restart and Run all your cells to see the correct output after you fix your answer for Question 14 (and Question 15).

```python
# compute and store the answer in the variable 'hitchcock_films', then display it

hitchcock_films = find_specific_movies(copy.deepcopy(movies), 'Alfred Hitchcock')
hitchcock_films
```

Output:
```
[{'title': 'The Wrong Man',
  'year': 1956,
  'duration': 105,
  'genres': ['Drama', 'Film-Noir'],
  'rating': 7.4,
  'directors': ['Alfred Hitchcock'],
  'cast': ['Henry Fonda', 'Vera Miles', 'Anthony Quayle', 'Harold J. Stone']},
 {'title': 'Lifeboat',
  'year': 1944,
  'duration': 97,
  'genres': ['Drama', 'War'],
  'rating': 7.6,
  'directors': ['Alfred Hitchcock'],
  'cast': ['Tallulah Bankhead',
   'John Hodiak',
   'Walter Slezak',
   'William Bendix']},
 {'title': 'Young and Innocent',
  'year': 1937,
  'duration': 80,
  'genres': ['Crime', 'Mystery', 'Romance'],
  'rating': 6.8,
  'directors': ['Alfred Hitchcock'],
  'cast': ['Nova Pilbeam',
   'Derrick De Marney',
   'Percy Marmont',
   'Edward Rigby']},
 {'title': 'Marnie',
  'year': 1964,
  'duration': 130,
  'genres': ['Crime', 'Drama', 'Mystery'],
  'rating': 7.1,
  'directors': ['Alfred Hitchcock'],
  'cast': ['Tippi Hedren', 'Sean Connery', 'Martin Gabel', 'Louise Latham']},
 {'title': 'I Confess',
  'year': 1953,
  'du
```

```python
grader.check("q14")
```

Output:
```
q14 results: All test cases passed!
q14 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`find_specific_movies` function is not used to answer (-4)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

**Question 15:** List all the movies that contain the string `"Once Upon a Time in"` in their `title`.

Your answer **must** be a **list** of **dictionaries**.

You **must** answer this question by calling `find_specific_movies` with the keyword `"Once Upon a Time in"`.

**Important Hint:**  If you did not notice the flaw with the definition of `find_specific_movies` before, you are likely to have run into an issue with this question. It is likely that you will see that your output for this question is an empty list. To see why this happened, find the value of `len(movies)` and see if it is equal to the value you found earlier.

```python
# compute and store the answer in the variable 'once_upon_a_time_in_movies', then display it
once_upon_a_time_in_movies = find_specific_movies(copy.deepcopy(movies), "Once Upon a Time in")
once_upon_a_time_in_movies
```

Output:
```
[{'title': 'Once Upon a Time in Mumbaai',
  'year': 2010,
  'duration': 134,
  'genres': ['Action', 'Crime', 'Drama'],
  'rating': 7.4,
  'directors': ['Milan Luthria'],
  'cast': ['Ajay Devgn', 'Emraan Hashmi', 'Kangana Ranaut', 'Prachi Desai']},
 {'title': 'Once Upon a Time in Deadwood',
  'year': 2019,
  'duration': 85,
  'genres': ['Action', 'Western'],
  'rating': 3.0,
  'directors': ['Rene Perez'],
  'cast': ['Robert Bronzi', 'Michael Paré', 'Karin Brauns', 'Lauren Compton']},
 {'title': 'Once Upon a Time in China II',
  'year': 1992,
  'duration': 113,
  'genres': ['Action', 'Adventure', 'Biography'],
  'rating': 7.3,
  'directors': ['Hark Tsui'],
  'cast': ['Jet Li', 'Rosamund Kwan', 'Siu Chung Mok', 'David Chiang']},
 {'title': 'Once Upon a Time in High School: The Spirit of Jeet Kune Do',
  'year': 2004,
  'duration': 116,
  'genres': ['Action', 'Drama', 'History'],
  'rating': 7.4,
  'directors': ['Ha Yoo'],
  'cast': ['Sang-woo Kwon', 'Lee Jung-Jin', 'Ga-in Han', 'Hyo-jun P
```

```python
grader.check("q15")
```

Output:
```
q15 results: All test cases passed!
q15 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`find_specific_movies` function is not used to answer (-4)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

### Function 5: `bucketize_by_genre(movies)`

We require you to complete the below function to answer the next several questions (this is a **requirement**, and you will **lose points** if you do not implement this function).

```python
def bucketize_by_genre(movies):
    x = {}
    for movie in movies:
        for genre in movie['genres']:
            if genre not in x:
                x[genre] = [movie]
            else:
                x[genre].append(movie)
    return x
    """bucketize_by_genre(movies) takes a list of movie dictionaries;
    it returns a dict in which each genre is a key and
    the value is a list of all movies that contain that genre"""
    # TODO: initialize a dictionary
    # TODO: loop through all movies
    # TODO:     loop through all genres in this movie
    # TODO:         if this genre is not a key in our dictionary, set the value associated with this genre to an empty list
    # TODO:         add the movie to the list associated with this genre
    # TODO: return the dictionary
```

```python
grader.check("bucketize_by_genre")
```

Output:
```
bucketize_by_genre results: All test cases passed!
bucketize_by_genre - 1 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'function logic is incorrect (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
bucketize_by_genre - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'function is called more than once with the same input list (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

**Warning:** You are **not** allowed to call `bucketize_by_genre` more than once on the full list of movies (`movies`) in your notebook. The Gradescope autograder will **deduct points** if you call `bucketize_by_genre` again on `movies`.

```python
# call the function bucketize_by_genre on 'movies' and store it in the variable 'genre_dict'
# do NOT display the output directly
x = copy.deepcopy(movies)
genre_dict = bucketize_by_genre(x)

```

```python
grader.check("genre_dict")
```

Output:
```
genre_dict results: All test cases passed!
genre_dict - 1 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'data structure is defined incorrectly (-3)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
genre_dict - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`bucketize_by_genre` function is not used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
genre_dict - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'data structure is redefined or modified after its initialization (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

**Question 16:** How many **unique** movie `genres` are present in the dataset?

```python
# compute and store the answer in the variable 'num_genres', then display it
num_genres = len(genre_dict)
num_genres
```

Output:
```
24
```

```python
grader.check("q16")
```

Output:
```
q16 results: All test cases passed!
q16 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q16 - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'answer uses loops to find number of unique movie genres (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q16 - 4 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`genre_dict` data structure is not used to answer (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

**Question 17:** How many *Romance* movies (i.e. movies with *Romance* as one of their `genres`) do we have in the dataset released **after** the `year` *2019*?

Your output **must** be an **int**. You **must** use the `genre_dict` data structure to answer this question.

```python
# compute and store the answer in the variable 'romance_after_2019', then display it
romance_after_2019 = genre_dict['Romance']
x = 0
for movie in romance_after_2019:
    if movie['year'] > 2019:
        x += 1
romance_after_2019 = x
romance_after_2019
```

Output:
```
1011
```

```python
grader.check("q17")
```

Output:
```
q17 results: All test cases passed!
q17 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q17 - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'answer uses loops to iterate over dictionary keys (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q17 - 4 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`genre_dict` data structure is not used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

**Question 18:** List the `title` of all *Documentary* movies (i.e. movies with *Documentary* as one of their `genres`) with `duration` **larger** than *240* in the dataset.

Your output **must** be a **list** of **strings**. You **must** use the `genre_dict` data structure to answer this question.

```python
# compute and store the answer in the variable 'docus_longer_than_240_mins', then display it
docus_longer_than_240_mins = genre_dict['Documentary']
x = []
for movie in docus_longer_than_240_mins:
    if movie['duration'] > 240:
        x.append(movie['title'])
docus_longer_than_240_mins = x
docus_longer_than_240_mins
```

Output:
```
['The Hour of the Furnaces',
 'September 11: The New Pearl Harbor',
 'The Cure for Insomnia',
 'At Berkeley',
 '16 Days of Glory',
 'Route One USA']
```

```python
grader.check("q18")
```

Output:
```
q18 results: All test cases passed!
q18 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'correct comparison operator is not used (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q18 - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q18 - 4 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'answer uses loops to iterate over dictionary keys (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q18 - 5 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`genre_dict` data struc
```

**Question 19:** Which movie `genre` is *Blake Lively* the most involved with?

We say that *Blake Lively* is involved with a movie if she is **either** a *director* **or** a *cast member* of that movie. You need to find the movie `genre` with the **highest** number of movies that *Blake Lively* is involved with. There is a **unique** `genre` that *Blake Lively* is most involved with. You do **not** have to worry about breaking ties.

You may **assume** that there are **no** movies in the dataset with the string `"Blake Lively"` in their `title`.

**Hint:** You can combine the *two* functions above to bucketize the movies that *Blake Lively* was involved with by their `genres`. Then, you can loop through each genre to find the one with the most number of movies in it.

```python
# compute and store the answer in the variable 'blake_lively_genre', then display it
blake_lively_genre = {}
x = 0
y = ''
for genre in genre_dict:
    for movie in genre_dict[genre]:
        if 'Blake Lively' in movie['cast']:
            if genre not in blake_lively_genre:
                blake_lively_genre[genre] = 1
            else:
                blake_lively_genre[genre] += 1

for key in blake_lively_genre:
    if blake_lively_genre[key] > x:
        x = blake_lively_genre[key]
        y = key

blake_lively_genre = y
blake_lively_genre
```

Output:
```
'Drama'
```

```python
grader.check("q19")
```

Output:
```
q19 results: All test cases passed!
q19 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'only movies that *Blake Lively* was `cast` in are considered (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q19 - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'movies other than those involving *Blake Lively* are considered (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
q19 - 4 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

**Question 20:** Who are the `directors` of the *Horror* movies with the **highest** `rating` in the movies dataset?

There are **multiple** *Horror* movies in the dataset with the joint highest rating. You **must** output a **list** of **strings** containing the **names** of **all** the `directors` of **all** these movies.

**Hint:** If you are unsure how to efficiently add the elements of one list to another, please review any of the lecture notes from the October 9 lectures ([Mike](https://canvas.wisc.edu/courses/374263/files/folder/Mikes_Lecture_Notes/Lec14_Lists), [Gurmail](https://git.doit.wisc.edu/cdis/cs/courses/cs220/cs220-lecture-material/-/tree/main/f23/Gurmail_Lecture_Notes/14_Lists), or [Cole](https://git.doit.wisc.edu/cdis/cs/courses/cs220/cs220-lecture-material/-/tree/main/f23/Cole_Lecture_Notes/14_Lists)).

```python
# compute and store the answer in the variable 'max_horror_rating_directors', then display it

max_horror_rating_directors = []
high_rating = 0
for movie in genre_dict['Horror']:
    if movie['rating'] > high_rating:
        high_rating = movie['rating']
        max_horror_rating_directors = copy.copy(movie['directors'])
    elif movie['rating'] == high_rating:
        max_horror_rating_directors.extend(movie['directors'])
max_horror_rating_directors
```

Output:
```
['Srini Hanumantharaju', 'Srini', 'Chris Delforce']
```

```python
grader.check("q20")
```

Output:
```
q20 results: All test cases passed!
q20 - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```

```python
grader.check("general_deductions")
```

Output:
```
general_deductions results: All test cases passed!
general_deductions - 1 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'Outputs not visible/did not save the notebook file prior to running the cell containing "export". We cannot see your output if you do not save before generating the zip file. (-3)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
general_deductions - 2 message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'Used concepts/modules such as csv.DictReader and pandas not covered in class yet. Note that built-in functions that you have been introduced to can be used. (-3)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
general_deductions - 3 message: Note that the Gradescope autograder will deduct points if your code fails the following rub
```

```python
grader.check("summary")
```

Output:
```
summary results: All test cases passed!
```

## Submission
It is recommended that at this stage, you Restart and Run all Cells in your notebook.
That will automatically save your work and generate a zip file for you to submit.

**SUBMISSION INSTRUCTIONS**:
1. **Upload** the zipfile to Gradescope.
2. If you completed the project with a **partner**, make sure to **add their name** by clicking "Add Group Member"
in Gradescope when uploading the zip file.
3. Check **Gradescope** results as soon as the auto-grader execution gets completed.
4. Your **final score** for this project is the score that you see on **Gradescope**.
5. You are **allowed** to resubmit on Gradescope as many times as you want to.
6. **Contact** a TA/PM if you lose any points on Gradescope for any **unclear reasons**.

```python
# running this cell will create a new save checkpoint for your notebook
from IPython.display import display, Javascript
display(Javascript('IPython.notebook.save_checkpoint();'))
```

```python
!jupytext --to py p8.ipynb
```

```python
public_tests.check_file_size("p8.ipynb")
grader.export(pdf=False, run_tests=False, files=["p8.py"])
```

Tests for bucketize_by_genre (points: 0)
```
>>> 
>>> public_tests.rubric_check('bucketize_by_genre: function logic is incorrect')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'function logic is incorrect (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('bucketize_by_genre: function is called more than once with the same input list')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'function is called more than once with the same input list (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for find_specific_movies (points: 0)
```
>>> 
>>> public_tests.rubric_check('find_specific_movies: function is redefined or defined more than once')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'function is redefined or defined more than once (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for general_deductions (points: 0)
```
>>> 
>>> public_tests.rubric_check('general_deductions: Outputs not visible/did not save the notebook file prior to running the cell containing "export". We cannot see your output if you do not save before generating the zip file.')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'Outputs not visible/did not save the notebook file prior to running the cell containing "export". We cannot see your output if you do not save before generating the zip file. (-3)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('general_deductions: Used concepts/modules such as csv.DictReader and pandas not covered in class yet. Note that built-in functions that you have been introduced to can be used.')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'Used concepts/modules such as csv.DictReader and pandas not covered in class yet. Note that built-in functions that you have been introduced to can be used. (-3)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('general_deductions: Large outputs such as movies are displayed in the notebook.')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'Large outputs such as movies are displayed in the notebook. (-3)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('general_deductions: Import statements are not mentioned in the required cell at the top of the notebook.')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'Import statements are not mentioned in the required cell at the top of the notebook. (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for genre_dict (points: 0)
```
>>> 
>>> public_tests.rubric_check('genre_dict: data structure is defined incorrectly')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'data structure is defined incorrectly (-3)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('genre_dict: `bucketize_by_genre` function is not used to answer', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`bucketize_by_genre` function is not used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('genre_dict: data structure is redefined or modified after its initialization')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'data structure is redefined or modified after its initialization (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for get_mapping (points: 0)
```
>>> 
>>> public_tests.rubric_check('get_mapping: function logic is incorrect')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'function logic is incorrect (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('get_mapping: function is defined more than once')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'function is defined more than once (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for get_movies (points: 0)
```
>>> 
>>> public_tests.rubric_check('get_movies: function is called more than once with the same dataset')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'function is called more than once with the same dataset (-4)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('get_movies: function logic is incorrect')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'function logic is incorrect (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('get_movies: `get_mapping` and `get_raw_movies` functions are not used', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`get_mapping` and `get_raw_movies` functions are not used (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('get_movies: function loops through IDs in the `mapping_path` file')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'function loops through IDs in the `mapping_path` file (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for get_raw_movies (points: 0)
```
>>> 
>>> public_tests.rubric_check('get_raw_movies: function logic is incorrect')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'function logic is incorrect (-3)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('get_raw_movies: column indices are hardcoded instead of using the header to identify column names')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'column indices are hardcoded instead of using the header to identify column names (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('get_raw_movies: function is defined more than once')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'function is defined more than once (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for movies (points: 0)
```
>>> 
>>> public_tests.rubric_check('movies: data structure is defined incorrectly')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'data structure is defined incorrectly (-3)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('movies: data structure is redefined or modified after its initialization')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'data structure is redefined or modified after its initialization (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for q1 (points: 0)
```
>>> public_tests.check('q1', small_mapping)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q1: `get_mapping` function is not used to answer', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`get_mapping` function is not used to answer (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for q10 (points: 0)
```
>>> public_tests.check('q10', second_movie_cast_small_movies)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q10: incorrect logic is used to answer')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q10: answer uses loops to iterate over dictionary keys')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'answer uses loops to iterate over dictionary keys (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q10: `small_movies` data structure is not used to answer', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`small_movies` data structure is not used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for q11 (points: 0)
```
>>> public_tests.check('q11', last_movie_directors_small_movies)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q11: incorrect logic is used to answer')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q11: answer uses loops to iterate over dictionary keys')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'answer uses loops to iterate over dictionary keys (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q11: `small_movies` data structure is not used to answer', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`small_movies` data structure is not used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for q12 (points: 0)
```
>>> public_tests.check('q12', movies_22202_22220)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q12: `movies` data structure is not used to answer', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`movies` data structure is not used to answer (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> public_tests.rubric_check('q12: public tests')
All test cases passed!
```


Tests for q13 (points: 0)
```
>>> public_tests.check('q13', num_movies_2023)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q13: incorrect logic is used to answer')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q13: answer uses loops to iterate over dictionary keys')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'answer uses loops to iterate over dictionary keys (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q13: `movies` data structure is not used to answer', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`movies` data structure is not used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for q14 (points: 0)
```
>>> public_tests.check('q14', hitchcock_films)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q14: `find_specific_movies` function is not used to answer', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`find_specific_movies` function is not used to answer (-4)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for q15 (points: 0)
```
>>> public_tests.check('q15', once_upon_a_time_in_movies)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q15: `find_specific_movies` function is not used to answer', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`find_specific_movies` function is not used to answer (-4)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for q16 (points: 0)
```
>>> public_tests.check('q16', num_genres)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q16: incorrect logic is used to answer')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q16: answer uses loops to find number of unique movie genres')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'answer uses loops to find number of unique movie genres (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q16: `genre_dict` data structure is not used to answer', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`genre_dict` data structure is not used to answer (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for q17 (points: 0)
```
>>> public_tests.check('q17', romance_after_2019)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q17: incorrect logic is used to answer')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q17: answer uses loops to iterate over dictionary keys')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'answer uses loops to iterate over dictionary keys (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q17: `genre_dict` data structure is not used to answer', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`genre_dict` data structure is not used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> public_tests.rubric_check('q17: public tests')
All test cases passed!
```


Tests for q18 (points: 0)
```
>>> public_tests.check('q18', docus_longer_than_240_mins)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q18: correct comparison operator is not used', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'correct comparison operator is not used (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q18: incorrect logic is used to answer')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q18: answer uses loops to iterate over dictionary keys')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'answer uses loops to iterate over dictionary keys (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q18: `genre_dict` data structure is not used to answer', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`genre_dict` data structure is not used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for q19 (points: 0)
```
>>> public_tests.check('q19', blake_lively_genre)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q19: only movies that *Blake Lively* was `cast` in are considered')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'only movies that *Blake Lively* was `cast` in are considered (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q19: movies other than those involving *Blake Lively* are considered')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'movies other than those involving *Blake Lively* are considered (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q19: incorrect logic is used to answer')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> public_tests.rubric_check('q19: public tests')
All test cases passed!
```


Tests for q2 (points: 0)
```
>>> public_tests.check('q2', nm1289434_value)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q2: `small_mapping` data structure is not used to answer', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`small_mapping` data structure is not used to answer (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q2: answer uses loops to iterate over dictionary keys')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'answer uses loops to iterate over dictionary keys (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for q20 (points: 0)
```
>>> public_tests.check('q20', max_horror_rating_directors)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q20: incorrect logic is used to answer')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> public_tests.rubric_check('q20: public tests')
All test cases passed!
```


Tests for q3 (points: 0)
```
>>> public_tests.check('q3', nm_values)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q3: answer does not check only for keys that begin with *nm*')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'answer does not check only for keys that begin with *nm* (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q3: `small_mapping` data structure is not used to answer', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`small_mapping` data structure is not used to answer (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for q4 (points: 0)
```
>>> public_tests.check('q4', nm_murphy)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q4: answer does not exclude people whose first name or middle name is Murphy')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'answer does not exclude people whose first name or middle name is Murphy (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q4: answer does not exclude people whose last name contains Murphy as a substring')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'answer does not exclude people whose last name contains Murphy as a substring (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q4: answer does not exclude movie titles')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'answer does not exclude movie titles (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q4: `small_mapping` data structure is not used to answer', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`small_mapping` data structure is not used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for q5 (points: 0)
```
>>> public_tests.check('q5', raw_small_movies)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q5: `get_raw_movies` function is not used to answer', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`get_raw_movies` function is not used to answer (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for q6 (points: 0)
```
>>> public_tests.check('q6', num_cast_first_movie)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q6: answer uses loops to find number of cast members')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'answer uses loops to find number of cast members (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q6: answer uses loops to iterate over dictionary keys')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'answer uses loops to iterate over dictionary keys (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q6: incorrect logic is used to answer')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q6: `raw_small_movies` data structure is not used to answer', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`raw_small_movies` data structure is not used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for q7 (points: 0)
```
>>> public_tests.check('q7', first_actor_id_first_movie)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q7: incorrect logic is used to answer')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q7: answer uses loops to iterate over dictionary keys')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'answer uses loops to iterate over dictionary keys (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q7: `raw_small_movies` data structure is not used to answer', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`raw_small_movies` data structure is not used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for q8 (points: 0)
```
>>> public_tests.check('q8', small_movies)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q8: `get_movies` function is not used to answer', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`get_movies` function is not used to answer (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for q9 (points: 0)
```
>>> public_tests.check('q9', second_movie_title_small_movies)
All test cases passed!
```
```
>>> 
>>> public_tests.rubric_check('q9: incorrect logic is used to answer')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'incorrect logic is used to answer (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q9: answer uses loops to iterate over dictionary keys')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'answer uses loops to iterate over dictionary keys (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('q9: `small_movies` data structure is not used to answer', False)
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - '`small_movies` data structure is not used to answer (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for small_movies (points: 0)
```
>>> 
>>> public_tests.rubric_check('small_movies: data structure is defined incorrectly')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'data structure is defined incorrectly (-2)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.
```
>>> 
>>> public_tests.rubric_check('small_movies: data structure is redefined or modified after its initialization')
All test cases passed!
```
success_message: Note that the Gradescope autograder will deduct points if your code fails the following rubric point - 'data structure is redefined or modified after its initialization (-1)'. The public tests cannot determine if your code satisfies these requirements. Verify your code manually.


Tests for summary (points: 127)
```
>>> public_tests.get_summary()
Total Score: 100/100
```
